```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

```python
from hashlib import sha1
```

# DSCI 521 - Lab 1


## Submission Instructions 
rubric={mechanics:4}

As you can see in the "rubric" line above, you can receive 4 marks for correctly submitting this assignment by following these instructions after answering all the questions:
    
- [Follow the general lab instructions](https://ubc-mds.github.io/resources_pages/general_lab_instructions/).
- [Click here to view a description of the rubrics used to grade the questions](https://github.com/UBC-MDS/public/tree/master/rubric).
- Push your `.ipynb` file to your GitHub repository for this lab.
- Submit your `.ipynb` file to Gradescope.
- Include a clickable link to your GitHub repo for the lab just below this cell (it should look something like this https://github.ubc.ca/MDS-2024-25/DSCI_521_labX_yourcwl.
- Don't forget to commit 3 times (covered in lecture 2) and also push your work to GitHub.


## 1 Warm up activity

First, let's become familiar with the vocabulary. 

```python
...

### You have to complete the exercise adding the number of definition that match. For example: Z = 23

A
B
C
D
E
F
G
H
I
J
K
L
M
N

```

```python
grader.check("q11")
```

<!-- BEGIN QUESTION -->

### Question 1.2
rubric={autograder:2}
    
Select the pairs that match. If you are unsure you can check [The Carpentries Glosario](https://glosario.carpentries.org/en/) to read the full description of each item.


|                              | Symbol              |
|------------------------------|---------------------|
| A. Root directory (see note) | 1. `$`              |
| B. Parent directory          | 2. `.`              |
| C. Current working directory | 3. `..`             |
| D. Home directory            | 4.  `/`             |
| E. Command line argument     | 5. `~`              |
| F. Prompt (R)                | 6.  `-y` or `--yes` |
| G. Prompt (Python)           | 7. `>>>`            |
| H. Prompt (Bash)             | 8. `>`              |

* Note: Notice that there are two meanings for the `/` character. When it appears inside a path, it’s just a separator.


```python
...

### You have to complete the exercise adding the number of definition that match. For example: Z = 23

A
B
C
D
E
F
G
H

```

```python
grader.check("q12")
```

<!-- END QUESTION -->

## 2. Using the command line to navigate your filesystem

Below are some questions to check-in on your understanding of the 3 key commands we need to navigate your filesystem (`cd`, `ls` and `pwd`), as well as some other concepts that are needed to succesfully do this.

<!-- BEGIN QUESTION -->

### Question 2.1
rubric={reasoning:1}

With your own words and in 1-2 sentences, describe what the `cd` command does.
If there is a default behaviour for this command (behaviour when no argument is given), also describe that.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.2
rubric={reasoning:1}

With your own words and in 1-2 sentences, describe what the `pwd` command does. If there is a default behaviour for this command (behaviour when no argument is given), also describe that.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.3
rubric={reasoning:1}

With your own words and in 1-2 sentences, describe what the `ls` command does. If there is a default behaviour for this command (behaviour when no argument is given), also describe that.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2.4
rubric={autograde:1}

If `pwd` results in `/Users/tiffany/data/`, which of the following commands could Tiffany use to navigate to her home directory, which is `/Users/tiffany`?

A. `cd .`

B. `cd /`

C. `cd tiffany`

D. `cd ..`

    
*Assign your answer to an object called `answer_14`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

```python
...
answer_14
```

```python
grader.check("q24")
```

### Question 2.5
rubric={autograde:1}

Using the filesystem diagram below, if `pwd` displays `/Users/backup/`, what will `ls ..` display?

A. `..: No such file or directory`

B. `thing/ backup/`

C. `Users/`

D. `original/ pnas_final/ pnas_sub/`
    
*Assign your answer to an object called `answer_15`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

![](imgs/filesystem-challenge.jpg)

```python
...
answer_15
```

```python
grader.check("q25")
```

<!-- BEGIN QUESTION -->

### Question 2.6
rubric={reasoning:2}
    
After running the command `pwd` we get the following output: `C:/Users/Florencia/dashboards`. Complete the table below with the equivalent command but using an absolute path. The first row is complete as an example:

| Relative path              | Absolute path                                 |
|----------------------------|-----------------------------------------------|
| `cd shiny-app`             | 1.  `cd /Users/Florencia/dashboard/shiny-app` |
| `cd ../../Varada/571-labs` | 2. `cd`                                       |
| `cd ..`                    | 3. `cd`                                       |
| `cd `                      | 4. `cd`                                       |
| `cd /`                     | 5. `cd`                                       |
| `cd ~/notes`               | 6. `cd`                                       |
| `cd ~/..`                  | 7. `cd`                                       |
| `cd ~/../../Data/mds_2022` | 8. `cd`                                       |

![](imgs/diagram-files-activity.png)


You can check [The Carpentries Glosario](https://glosario.carpentries.org/en/) to learn about more terms or to read the full description of each item.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Before starting: setting your SSH keys
rubric={mechanics:2}

Follow the instructions on [how to setup SSH authentication with GitHub for your operating system](https://docs.github.com/en/free-pro-team@latest/github/authenticating-to-github/generating-a-new-ssh-key-and-adding-it-to-the-ssh-agent).

When you have setup your keys, run ssh -T `git@github.ubc.ca` from a terminal. If the message says that you have successfully authenticated take a screenshot of your terminal showing both the input and the output and paste it below this cell. It should look similar to this:

![](imgs/ssh-login.png)

Type `git remote -v` in your local repo folder and take note of what URLs. Next, create a new folder on your computer and clone your lab repo into this folder using the SSH instead of HTTPS. In your new folder, type `git remote -v`. Paste the URLs of the SSH and HTTPS remotes below and explain how they differ.


## 3. Git and GitHub

In class we introduced Git and GitHub to access, log changes, and submit your homework. Below are some multiple choice questions to check-in on your understanding of the key commands needed to do this (`clone`, `add`, `commit`, `push` and `pull`), as well other key concepts.

### Question 3.1
rubric={autograde:6}

Now you know the basics of the Git/GitHub workflows!

In this assignment you are going to create a repository and interact with it.
We suggest you to do so to practice the steps at the same time you answer the questions.

First, let's to match some of the new terms we have learned through this chapter with a related concept/word.
Assign the variable the corresponding lower-case letter.

|                                   | related concept / word |
|-----------------------------------|------------------------|
| one. Git repository               | a. The staging area    |
| two. Remote repository location   | b. `.git` folder       |
| three. Git add                    | c. `git pull`          |
| four. Clone a repository          | d. GitHub              |
| five. Update the local repository | e. SSH keys            |

Complete the following sentences selecting one of the following options.
Assign the variable the provided string.

- **If you are working on the remote you are using [six].**
    - git
    - github
    - jupyterlab
    - the terminal
<br /><br />

- **One of the benefits of using version control is that you can [seven].**
    - open your files using VS Code
    - have a back up of your code
    - track the history of your project
<br /><br />

- **The expression "[eight] to GitHub" is referring to the step when you update the GitHub repository with the local changes.**
    - pulling
    - pushing
    - cloning 
<br /><br />

- **You record a new version of your project each time you [nine] your changes.**
    - commit
    - stage
    - push
    - clone
<br /><br />

- **The HASH is a string of characters that serves as [ten].**
    - commit identifier
    - Git repository identifier
    - commit message
<br /><br />

- **If you create a new file into a Git repository, its status will be [eleven].**
    - changed
    - untracked
    - staged
    - committed
<br /><br />

- **If you want to check how many files do you have in the staging area you have to run the command [twelve].**
    - git add
    - git status
    - git log
    - git commit
<br /><br />

*Assign your answer to an object called `one`. Make sure your answer is a lowercase letter/word and is surrounded by quotation marks (e.g. `"b"` or `"git pull"`).

```python
...

one
two
three
four
five
six
seven
eight
nine
ten
eleven
twelve

```

```python
grader.check("q31")
```

### Question 3.2
rubric={reasoning:1}

With your own words and in 1-2 sentences, explain why you should use `git clone` to get a copy of your homework (or other Git repositories) from GitHub as opposed to clicking on the "Download ZIP" link.

### Question 3.3
rubric={autograde:1}

Which command(s) below would save the changes of `README.md` to my local Git repository?
 
A.<br>
`git init README.md`<br>
`git commit -m "Add authors and copyright"`<br><br>

B.<br>
`git commit -m "Add authors and copyright"`<br><br>

C.<br>
`git add README.md`<br>
`git commit -m "Add authors and copyright"`<br><br>

D.<br>
`git commit -m README.md "Add authors and copyright"`<br>

*Assign your answer to an object called `answer_33`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).* 

```python
...
answer_33
```

```python
grader.check("q33")
```

### Question 3.4
rubric={reasoning:1}

With your own words and in 1-2 sentences, compare and contrast `git pull` and `git push`.

### Question 3.5
rubric={autograde:1}

When working on a project that is under version control (*e.g.,* MDS homework/labs) you should `commit` to you local Git repository:

A. after every new unit/addition/fix you add to the project

B. 2-3 times

C. only once, shortly before the deadline

D. obsessively, like every 5 minutes
    
*Assign your answer to an object called `answer_35`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

```python
...
answer_35
```

```python
grader.check("q35")
```

### Question 3.6
rubric={autograde:1}

When working on a project that is under version control (*e.g.,* MDS homework/labs) you should `push` to online GitHub repository:

A. obsessively, like every 5 minutes

B. at least at the end of every session you work on the project

C. 2-3 times

D. only once, shortly before the deadline
    
*Assign your answer to an object called `answer_36`. Make sure your answer is an uppercase letter and is surrounded by quotation marks (e.g. `"F"`).*

```python
...
answer_36
```

```python
grader.check("q36")
```

<!-- BEGIN QUESTION -->

### Question 3.7

#### Extra: Git + Globbing
rubric={reasoning:4.5}
    
In the past chapter we have learn how to use the shell. If you are using Git through the terminal you can benefit from from globbing and some of the operators we learned there.

Let's say that you have this files with changes on your GitHub repository:

```
2022-03-02_vancouver_site1.csv
2022-03-04_vancouver_site2.csv
2022-05-20_okenagan_site55.csv
```

but you only want to commit the changes in the files that say "Vancouver" as you are still waiting information to complete the last file with the data from Okenagan.

Select the correct answers (more than one option could be correct)

1 - How would you add to the staging area all the files including the word "vancouver" using the command line? Select all the correct answers.

A. `git add .`  
B. `git add 2022-03-02_vancouver_site1.csv 2022-03-02_vancouver_site2.csv`  
C. `git add *vancouver*`  
D.   
`git add 2022-03-02_vancouver_site1.csv`    
`git add 2022-03-02_vancouver_site2.csv`  

2 - From all the correct answers, which one you consider it would be more appropiate to use if instead of 3 files you would have 500 files?

YOUR ANSWER GOES IN THIS CELL

(paste your log as per the instructions above)

<!-- END QUESTION -->

---

# Submission 

When you are ready to submit your assignment do the following:

1. Run all cells in your notebook to make sure there are no errors by doing `Kernel -> Restart Kernel and Run All Cells...`
2. Enter Canvas and upload the assignment using Gradescope's drag and drop tool.
3. Don't forget to commit 3 times and also push your work to GitHub.

Remember that you will find a detailed set of instructions including videos on [Canvas](https://canvas.ubc.ca/courses/102033).

### Attribution

- [Data 8: The Foundations of Data Science](http://data8.org/), specifically Lab 01
- [Software Carpentry](https://software-carpentry.org/), the Unix Shell lesson.
- [Library Carpentry: Introduction to Working with Data (Regular Expressions)](https://librarycarpentry.org/lc-data-intro/)
- [Glossary - The Carpentries](https://glosario.carpentries.org/en/)

Tests for q11 (points: 3.5)
```
>>> assert sha1(str(A).encode('utf8')).hexdigest() == 'fe5dbbcea5ce7e2988b8c69bcfdfde8904aabc1f', 'A is the definition 8'
```
```
>>> assert sha1(str(B).encode('utf8')).hexdigest() == '17ba0791499db908433b80f37c5fbc89b870084b', 'B is the definition 11'
```
```
>>> assert sha1(str(C).encode('utf8')).hexdigest() == '356a192b7913b04c54574d18c28d46e6395428ab', 'C is the definition 1'
```
```
>>> assert sha1(str(D).encode('utf8')).hexdigest() == 'c1dfd96eea8cc2b62785275bca38ac261256e278', 'D is the definition 6'
```
```
>>> assert sha1(str(E).encode('utf8')).hexdigest() == '902ba3cda1883801594b6e1b452790cc53948fda', 'E is the definition 7'
```
```
>>> assert sha1(str(F).encode('utf8')).hexdigest() == 'bd307a3ec329e10a2cff8fb87480823da114f8f4', 'F is the definition 13'
```
```
>>> assert sha1(str(G).encode('utf8')).hexdigest() == '7b52009b64fd0a2a49e6d8a939753077792b0554', 'G is the definition 12'
```
```
>>> assert sha1(str(H).encode('utf8')).hexdigest() == '1b6453892473a467d07372d45eb05abc2031647a', 'H is the definition 4'
```
```
>>> assert sha1(str(I).encode('utf8')).hexdigest() == 'fa35e192121eabf3dabf9f5ea6abdbcbc107ac3b', 'I is the definition 14'
```
```
>>> assert sha1(str(J).encode('utf8')).hexdigest() == 'b1d5781111d84f7b3fe45a0852e59758cd7a87e5', 'J is the definition 10'
```
```
>>> assert sha1(str(K).encode('utf8')).hexdigest() == 'ac3478d69a3c81fa62e60f5c3696165a4e5e6ac4', 'K is the definition 5'
```
```
>>> assert sha1(str(L).encode('utf8')).hexdigest() == 'da4b9237bacccdf19c0760cab7aec4a8359010b0', 'L is the definition 2'
```
```
>>> assert sha1(str(M).encode('utf8')).hexdigest() == '77de68daecd823babbb58edb1c8e14d7106e83bb', 'M is the definition 3'
```
```
>>> assert sha1(str(N).encode('utf8')).hexdigest() == '0ade7c2cf97f75d009975f4d720d1fa6c19f4897', 'N is the definition 9'
```


Tests for q12 (points: None)
```
>>> assert sha1(str(A).encode('utf8')).hexdigest() == '1b6453892473a467d07372d45eb05abc2031647a', 'A is the definition 4'
```
```
>>> assert sha1(str(B).encode('utf8')).hexdigest() == '77de68daecd823babbb58edb1c8e14d7106e83bb', 'B is the definition 3'
```
```
>>> assert sha1(str(C).encode('utf8')).hexdigest() == 'da4b9237bacccdf19c0760cab7aec4a8359010b0', 'C is the definition 2'
```
```
>>> assert sha1(str(D).encode('utf8')).hexdigest() == 'ac3478d69a3c81fa62e60f5c3696165a4e5e6ac4', 'D is the definition 5'
```
```
>>> assert sha1(str(E).encode('utf8')).hexdigest() == 'c1dfd96eea8cc2b62785275bca38ac261256e278', 'E is the definition 6'
```
```
>>> assert sha1(str(F).encode('utf8')).hexdigest() == 'fe5dbbcea5ce7e2988b8c69bcfdfde8904aabc1f', 'F is the definition 8'
```
```
>>> assert sha1(str(G).encode('utf8')).hexdigest() == '902ba3cda1883801594b6e1b452790cc53948fda', 'G is the definition 7'
```
```
>>> assert sha1(str(H).encode('utf8')).hexdigest() == '356a192b7913b04c54574d18c28d46e6395428ab', 'H is the definition 1'
```


Tests for q24 (points: 1)
```
>>> assert sha1(str(answer_14).encode('utf8')).hexdigest() == '50c9e8d5fc98727b4bbc93cf5d64a68db647f04f', 'Incorrect answer, try again!'
```


Tests for q25 (points: 1)
```
>>> assert sha1(str(answer_15).encode('utf8')).hexdigest() == 'ae4f281df5a5d0ff3cad6371f76d5c29b6d953ec', 'Incorrect answer, try again!'
```


Tests for q31 (points: 6)
```
>>> assert sha1(str(one).encode('utf8')).hexdigest() == 'e9d71f5ee7c92d6dc9e92ffdad17b8bd49418f98', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(two).encode('utf8')).hexdigest() == '3c363836cf4e16666669a25da280a1865c2d2874', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(three).encode('utf8')).hexdigest() == '86f7e437faa5a7fce15d1ddcb9eaeaea377667b8', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(four).encode('utf8')).hexdigest() == '58e6b3a414a1e090dfc6029add0f3555ccba127f', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(five).encode('utf8')).hexdigest() == '84a516841ba77a5b4648de2cd0dfcb30ea46dbb4', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(six).encode('utf8')).hexdigest() == '64b2b6d12bfe4baae7dad3d018f8cbf6b0e7a044', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(seven).encode('utf8')).hexdigest() == '36cab06ad781f49e411f1c9466d9084195d9e879', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(eight).encode('utf8')).hexdigest() == 'adb1108fc750116ee99e4222566325dd7775e96e', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(nine).encode('utf8')).hexdigest() == '4015b57a143aec5156fd1444a017a32137a3fd0f', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(ten).encode('utf8')).hexdigest() == 'b1939f502fce507a9616e97bd8749fed83da3b7a', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(eleven).encode('utf8')).hexdigest() == 'd7390aa3e7289506746a90490cae21f0dfa03b28', 'Incorrect answer, try again!'
```
```
>>> assert sha1(str(twelve).encode('utf8')).hexdigest() == 'f2a29f2b358f8a8bc56896bdca251ef93927330b', 'Incorrect answer, try again!'
```


Tests for q33 (points: 1)
```
>>> assert sha1(str(answer_33).encode('utf8')).hexdigest() == '32096c2e0eff33d844ee6d675407ace18289357d', 'Incorrect answer, try again!'
```


Tests for q35 (points: 1)
```
>>> assert sha1(str(answer_35).encode('utf8')).hexdigest() == '6dcd4ce23d88e2ee9568ba546c007c63d9131c1b', 'Incorrect answer, try again!'
```


Tests for q36 (points: 1)
```
>>> assert sha1(str(answer_36).encode('utf8')).hexdigest() == 'ae4f281df5a5d0ff3cad6371f76d5c29b6d953ec', 'Incorrect answer, try again!'
```
